```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02A.ipynb")
```

# 🍕 Homework 2A: Food Safety

## 🐼 Cleaning and Exploring Data with Pandas

## ‼️ Due Date: Friday, February 6th, 11:59 PM
You must submit this assignment to Pensieve by the on-time deadline, Friday, February 6th, 11:59 PM. Please read the syllabus for the Slip Day policy. No late submissions beyond the details in the Slip Day policy will be accepted. While course staff is happy to help you if you encounter difficulties with submission, we may not be able to respond to late-night requests for assistance (TAs need to sleep, after all!). **We strongly encourage you to plan to submit your work to Pensieve several hours before the stated deadline.** This way, you will have ample time to contact staff for submission support. 

Please read the instructions carefully when you are submitting your work to Pensieve.

## 💪 Collaboration Policy

Data science is a collaborative activity. While you may talk with others about
the homework, we ask that you **write your solutions individually**. If you do
discuss the assignments with others, please **include their names** below.

**Collaborators**: *list collaborators here*


## 📜 This Assignment

In this homework, we will investigate restaurant food safety scores for restaurants in San Francisco. The scores and violation information have been [made available by the San Francisco Department of Public Health](https://data.sfgov.org/Health-and-Social-Services/Restaurant-Scores-LIVES-Standard/pyih-qa8i). The main goal for this assignment is to walk through the process of Data Cleaning and familiarise yourself with some of the `pandas` functions discussed in Pandas I and II. 


After this homework, you should be comfortable with:
* Reading CSV files, 
* Reading `pandas` documentation and using `pandas`,
* Working with data at different levels of granularity, and
* Identifying the type of data collected, missing values, anomalies, etc., and doing some basic analysis.

## 💯 Score Breakdown 
Question | Manual | Points
--- | --- | ---
1a | no | 1
1b | no | 2
1c | no | 1
2a | no | 1
2bi | no | 1
2bii | no | 1
2ci | no | 1
2cii | no | 1
2d | yes | 1
2e | no | 2
2f | no | 1
3a | no | 1
3b | no | 2
3ci | no | 1
3cii | no | 1
3ciii | no | 1
3civ | no | 1
4a | no | 1
4b | no | 1
4c | no | 1
5a | no | 2
5b | no | 2
Total | 1 | 27


## ✊ Before You Start

For each question in the assignment, please write down your answer in the answer cell(s) right below the question. 

We understand that it is helpful to have extra cells breaking down the process towards reaching your final answer. If you happen to create new cells below your answer to run codes, **NEVER** add cells between a question cell and the answer cell below it. It will cause errors when we run the autograder, and it will sometimes cause a failure to generate the PDF file.

**Important note: The local autograder tests will not be comprehensive. You can pass the automated tests in your notebook but still fail tests in the autograder.** Please be sure to check your results carefully.

Finally, unless we state otherwise, **do not use for loops or list comprehensions**. The majority of this assignment can be done using built-in commands in `pandas` and `NumPy`. Please also don't hard code anything. Our autograder isn't smart enough to check, but you're depriving yourself of key learning objectives if you write loops or comprehensions, and you also won't be ready for the midterm.

The cell below imports all the necessary libraries you need to use during this homework. Without running this cell, you will not be able to call the various `NumPy` and `pandas` functions we use later on, so please make sure you run it before starting to work on the homework.

```python
# DO NOT EDIT THIS CELL
import numpy as np
import pandas as pd
pd.options.mode.chained_assignment = None

import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
sns.set()
plt.style.use('fivethirtyeight')

from IPython.display import display, Image 
def display_figure_for_grader(fig):
    plotly.io.write_image(fig, 'temp.png')
    display(Image('temp.png'))    
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# 🚀 0: Obtaining the Data

## File Systems and I/O


In general, we will focus on using Python commands to investigate files. However, it can sometimes be easier to use shell commands in your local operating system. The following cells demonstrate how to do this.

```python
from pathlib import Path
data_dir = Path('.') # Create a Path object to the current working directory
data_dir.mkdir(exist_ok = True) # Create a directory 
file_path = data_dir / Path('data.zip') # Create a path to the .zip file that contains our datasets
dest_path = file_path
```

After running the cell above, if you list the contents of the directory containing this notebook, you should see `data.zip`.

*Note*: The command below starts with an `!`. This tells our Jupyter Notebook to pass this command to the operating system. In this case, the command is the `ls` Unix command, which lists files in the current directory.

```python
!ls 
```

## Loading Food Safety Data

We have data, but we don't have any specific questions about the data yet. Let's focus on understanding the structure of the data; this involves answering questions such as:

* Is the data in a standard format or encoding?
* Is the data organized in records?
* What are the fields in each record? (We sometimes also use the term 'feature' or 'attribute' as well, depending on the context)

Let's start by looking at the contents of `data.zip`. It's not just a single file but rather a compressed directory of multiple files. We could inspect it by uncompressing it using a shell command such as `!unzip data.zip`, but in this homework, we're going to do almost everything in Python for maximum portability.

## Looking Inside and Extracting the Zip Files

The following code blocks are for setup. Simply run the cells; **do not modify them**. Question 1a is where you will start to write code.

Here, we assign `my_zip` to a `zipfile.Zipfile` object representing `data.zip`, and assign `list_names` to a list of all the names of the contents in `data.zip`.

```python
import zipfile
my_zip = zipfile.ZipFile(dest_path, 'r') # Read .zip file at dest_path, which is data.zip
list_names = my_zip.namelist() # Create list of names of contents in data.zip
list_names
```

You may notice that we did not write `zipfile.ZipFile('data.zip', ...)`. Instead, we used `zipfile.ZipFile(dest_path, ...)`. In general, we **strongly suggest having your filenames hard coded as string literals only once** in a notebook. It is very dangerous to hardcode things twice because if you change one but forget to change the other, you can end up with bugs that are very hard to find.

Now, we display the files' names and their sizes.

```python
my_zip = zipfile.ZipFile(dest_path, 'r') # Read .zip file at dest_path, which is data.zip
for info in my_zip.infolist(): # Extract information about contents of .zip file, iterate through them
    print('{}\t{}'.format(info.filename, info.file_size)) # Print file name and uncompressed file size
```

Often when working with zipped data, we'll never unzip the actual zip file. This saves space on our local computer. However, for this homework the files are small, so we're just going to unzip everything. This has the added benefit that you can look inside the CSV files using a text editor, which might be handy for understanding the structure of the files. The cell below will unzip the CSV files into a sub-directory called `data`.

```python
data_dir = Path('.') # Create a path to current working directory
my_zip.extractall(data_dir) # Extract all contents from data.zip
!ls {data_dir / Path("data")} # List out contents of "data" sub-directory
```

The cell above created a folder called `data`, and in it, there should be five CSV files. Let's open up `legend.csv` to see its contents. To do this, click on the file icon on the top left to show the folders and files within the hw02A folder, then click `legend.csv`. The file will open up in another tab. You should see something that looks like this:

    "Minimum_Score","Maximum_Score","Description"
    0,70,"Poor"
    71,85,"Needs Improvement"
    86,90,"Adequate"
    91,100,"Good"

The `legend.csv` file does indeed look like a well-formed CSV file. Let's check the other three files. Rather than opening up each file manually, let's use Python to print out the first 5 lines of each. We defined a helper function for you that will allow you to retrieve the first N lines of a file as a list. For example, `head('data/legend.csv', 5)` will return the first 5 lines of `data/legend.csv.gz`. Run the cell below to print out the first 5 lines of all six files that we just extracted from the zip file.

```python
import os

def head(filename, lines=5):
    """
    Returns the first few lines of a file.
    
    filename: the name of the file to open
    lines: the number of lines to include
    
    return: A list of the first few lines from the file.
    """
    from itertools import islice
    with open(filename, "r") as f: # Open file for reading
        return list(islice(f, lines)) # Return first five lines of the file as a list

data_dir = "./" # Path to current working directory
for f in list_names: # Iterate through names of contents in "data" sub-directory
    if not os.path.isdir(f): # Check if specified path is a directory
        print(head(data_dir + f, 5), "\n") # Create path to respective file in "data" sub-directory, use `head` function defined above
```

## Reading in and Verifying Data

Based on the above information, let's attempt to load `bus.csv`, `ins2vio.csv`, `ins.csv`, and `vio.csv` into `pandas` `DataFrame`s with the following names: `bus`, `ins2vio`, `ins`, and `vio`, respectively.

*Note:* Because of character encoding issues, one of the files (`bus`) will require an additional argument `encoding='ISO-8859-1'` when calling `pd.read_csv`. At some point in your future, you should read all about [character encodings](https://diveintopython3.net/strings.html). We won't discuss these in detail in Data 100.

```python
# Path to the directory containing data
dsDir = Path('data')

bus = pd.read_csv(dsDir/'bus.csv', encoding='ISO-8859-1')
ins2vio = pd.read_csv(dsDir/'ins2vio.csv')
ins = pd.read_csv(dsDir/'ins.csv')
vio = pd.read_csv(dsDir/'vio.csv')

# This code is essential for the autograder to function properly. Do not edit
bus.loc[:, ["latitude", "longitude"]] = bus.loc[:, ["latitude", "longitude"]].apply(lambda x: np.round(x, 2))
ins_test = ins
```

Now that you've read the files let's try some `pd.DataFrame` methods ([docs](https://pandas.pydata.org/pandas-docs/version/1.4.3/reference/api/pandas.DataFrame.html)).
Use the `DataFrame.head` method to show the top few lines of the `bus`, `ins`, and `vio` `DataFrame`s. For example, running the cell below will display the first few lines of the `bus` `DataFrame`. 

```python
bus.head()
```

To show multiple return outputs in one single cell, you can use `display()`. 

```python
display(bus.head())
display(ins.head())
```

The `DataFrame.describe` method can also be handy for computing summaries of numeric columns of our `DataFrame`s. Try it out with each of our 4 `DataFrame`s. Below, we have used the method to give a summary of the `bus` `DataFrame`. 

```python
bus.describe()
```

Now, we perform some sanity checks for you to verify that the data was loaded with the correct structure.

First, we check the basic structure of the `DataFrame`s you created:

```python
assert all(bus.columns == ['business id column', 'name', 'address', 'city', 'state', 'postal_code',
                           'latitude', 'longitude', 'phone_number']) # Check that the `bus` DataFrame contains the right columns in the correct order
assert 6250 <= len(bus) <= 6260 # Check that the `bus` DataFrame has the correct number of rows

assert all(ins.columns == ['iid', 'date', 'score', 'type']) # Similar to above, check that the `ins` DataFrame contains the right columns in order
assert 26660 <= len(ins) <= 26670 # Similar to above, check that the `ins` DataFrame has the correct number of rows

assert all(vio.columns == ['description', 'risk_category', 'vid'])
assert 60 <= len(vio) <= 65

assert all(ins2vio.columns == ['iid', 'vid'])
assert 40210 <= len(ins2vio) <= 40220
```

*Fun Fact*: The series of `assert` statements above is similar to how course staff implement test cases for assignments! Next time you fail a public test case on an assignment, take a look at the error message and you will probably find the `assert` statement which is causing you to fail the test case.

Next we'll check that the statistics match what we expect. The following are hard-coded statistical summaries of the correct data.

```python
bus_summary = pd.DataFrame(**{'columns': ['business id column', 'latitude', 'longitude'],
 'data': {'business id column': {'50%': 75685.0, 'max': 102705.0, 'min': 19.0},
  'latitude': {'50%': -9999.0, 'max': 37.82, 'min': -9999.0},
  'longitude': {'50%': -9999.0,
   'max': 0.0,
   'min': -9999.0}},
 'index': ['min', '50%', 'max']})

ins_summary = pd.DataFrame(**{'columns': ['score'],
 'data': {'score': {'50%': 76.0, 'max': 100.0, 'min': -1.0}},
 'index': ['min', '50%', 'max']})

vio_summary = pd.DataFrame(**{'columns': ['vid'],
 'data': {'vid': {'50%': 103135.0, 'max': 103177.0, 'min': 103102.0}},
 'index': ['min', '50%', 'max']})

from IPython.display import display

print('What we expect from your Businesses DataFrame:')
display(bus_summary)
print('What we expect from your Inspections DataFrame:')
display(ins_summary)
print('What we expect from your Violations DataFrame:')
display(vio_summary)
```

The code below defines a testing function that we'll use to verify that your data has the same statistics as what we expect. Run these cells to define the function. The `df_allclose` function has this name because we are verifying that all of the statistics for your `DataFrame` are close to the expected values. Why not `df_allequal`? It's a bad idea in almost all cases to compare two floating point values like 37.780435, as rounding errors can cause spurious failures. Run the following cells to load some basic utilities (you do not need to change these at all):

```python
"""Run this cell to load this utility comparison function that we will use in various
tests below (both tests you can see and those we run internally for grading).

Do not modify the function in any way.
"""


def df_allclose(actual, desired, columns=None, rtol=5e-2):
    """Compare selected columns of two Dataframes on a few summary statistics.
    
    Compute the min, median, and max of the two Dataframes on the given columns, and compare
    that they match numerically to the given relative tolerance.
    
    If they don't match, an AssertionError is raised (by `numpy.testing`).
    """    
    # Summary statistics to compare on
    stats = ['min', '50%', 'max']
    
    # For the desired values, we can provide a full DF with the same structure as
    # the actual data or pre-computed summary statistics.
    # We assume a pre-computed summary was provided if a column is None. In that case, 
    # `desired` *must* have the same structure as the actual's summary
    if columns is None:
        des = desired
        columns = desired.columns
    else:
        des = desired[columns].describe().loc[stats]

    # Extract summary stats from actual DF
    act = actual[columns].describe().loc[stats]

    return np.allclose(act, des, rtol)
```

We will now explore each file in turn, including determining its granularity and exploring many of the variables individually. Let's begin with the businesses file, which has been read into the `bus` `DataFrame`.

<br/>

---

<br/>

# 📚 1: Examining the Business Data File

## 🚌 Question 1a

From its name alone, we expect the `bus.csv` file to contain information about the restaurants. Let's investigate the granularity of this dataset.

```python
bus.head()
```

The `bus` `DataFrame` contains a column called `business id column`, which probably corresponds to a unique business id.  However, we will first rename that column to `bid` for simplicity.

**Note**: In practice, we might want to do this renaming when the table is loaded, but for grading purposes, we will do it here.


```python
bus = bus.rename(columns={"business id column": "bid"})
```

Examining the entries in `bus`, is the `bid` unique for each record (i.e., each row of data)? Your code should compute the answer, i.e., don't just hard code `True` or `False`.

**Hint**: Use `value_counts()` (documentation [here](https://pandas.pydata.org/docs/reference/api/pandas.Series.value_counts.html)) or `unique()` (documentation [here](https://pandas.pydata.org/docs/reference/api/pandas.Series.unique.html)) to determine if the `bid` series has any duplicates.

```python
is_bid_unique = ...

is_bid_unique
```

```python
grader.check("q1a")
```

<br/>

---


## 🚌 Question 1b

We will now work with some important fields in `bus`.

1. Assign `top_names` to a `NumPy` array or list containing the top 7 most frequently used business names, from most frequent to least frequent. 
2. Assign `top_addresses` to a `NumPy` array or list containing the top 7 addresses where businesses are located, from most popular to least popular.

**Hint 1**: You may find `pd.Series.value_counts` helpful (documentation [here](https://pandas.pydata.org/docs/reference/api/pandas.Series.value_counts.html)).

**Hint 2**: You'll need to get the names and addresses, NOT the counts associated with each. Some way to **reset the index** would come in handy. If you're unsure how to do this, try looking through the class notes or using a search engine. Part of the goal of this course is to develop independent thinking in the context of the data science lifecycle, which can involve a fair bit of exploring and reading documentation. It may be a bit annoying at first, but you’ll get the hang of it. We’re here to help guide you on that path if you ever need support!

**Hint 3**: To check your answer, `top_names[0]` should return the string `Peet's Coffee & Tea`. It should not be a number.

```python
top_names = ...
top_addresses = ...

display(top_names)
display(top_addresses)
```

```python
grader.check("q1b")
```

<br/>

---


## 🚌 Question 1c

Based on the above exploration, what does each record in the `bus` `DataFrame` represent?

**A**. A city block.

**B**. A chain of restaurants.

**C**. One location of a restaurant.

**D**. A postal code.


Answer in the following cell. Your answer should be a string, either `"A"`, `"B"`, `"C"`, or `"D"`.


```python
q1c = ...
```

```python
grader.check("q1c")
```

<br/>

---

<br/>

# 🧹 2: Cleaning the Business Data Postal Codes

The business data contains postal code information that we can use to aggregate the ratings over regions of the city. Let's examine and clean the postal code field. The postal code (sometimes also called a [ZIP code](https://en.wikipedia.org/wiki/ZIP_Code)) partitions the city into regions:

<img src="https://gisgeography.com/wp-content/uploads/2023/07/San-Francisco-Zip-Code-Map-1-2048x2048.jpg" alt="ZIP Code Map" style="width: 600px">

<br/>

---


## 🥡 Question 2a

How many restaurants are in each ZIP code? 

In the cell below, create a **Series** where the index is the postal code, and the value is the number of records with that postal code. The `Series` should be in descending order of count. Do you notice any odd/invalid ZIP codes?


<!-- grade_id: cell-d2151d673e6c36a1, points: 0, grade: false, solution: true -->
```python
zip_counts = ...
print(zip_counts.to_string())
```

```python
grader.check("q2a")
```

<br/>

--- 

## 🥡 Question 2b

In Question 2a, we noticed a large number of potentially invalid ZIP codes (e.g., "Ca"). These are likely due to data entry errors. To get a better understanding of the potential errors in the zip codes, let's break down the problem into two parts.

### 🚀 Part I

Import a list of valid San Francisco ZIP codes by using `pd.read_json` to load the file `data/sf_zipcodes.json`, and store them as a **Series** in `valid_zip_codes`. As you may expect, `pd.read_json` works similarly to `pd.read_csv` but for JSON files (a different file format you'll learn more about in HW 3) that you can read more about [here](https://pandas.pydata.org/docs/reference/api/pandas.read_json.html). **Make sure that the resulting series (not a data frame!) is zero-indexed, with each value being a unique zip code.** If you are unsure of what data type a variable is, remember you can do `type(some_var_name)` to check!

```python
valid_zip_codes = ...
valid_zip_codes
```

```python
grader.check("q2bi")
```

Observe that `pd.read_json` reads data as integers by default. This isn't quite what we want! We would like to store ZIP codes as strings (you'll learn more about why soon!). To do that, we can use the `astype` function to generate a copy of the `pandas` `Series` stored as strings instead.

```python
valid_zip_codes = valid_zip_codes.astype("string")
```

If you're ever unsure about the data type of a variable, remember you can always check using the `type` function like below:

```python
type(valid_zip_codes.dtype)
```

Now it's time to do Part II. You will probably want to use the `Series.isin` function. For more information on this function see the [documentation](https://pandas.pydata.org/docs/reference/api/pandas.Series.isin.html). 

**Note:** You are welcome and, in fact, encouraged to search and read the documentation on the internet to complete the assignments in the course, even if the documentation is not linked explicitly.

### 🚀 Part II

 Construct a `DataFrame` containing only the businesses that **DO NOT** have valid ZIP codes.

```python
...
invalid_zip_bus = ...
invalid_zip_bus.head(20)
```

```python
grader.check("q2bii")
```

<br/>

--- 

## ❓ Question 2c

In the previous question, many of the businesses had a common invalid postal code that was likely used to encode a MISSING postal code. Do they all share a potentially "interesting address"? For that purpose, in the following cells, we will construct a series that counts the number of businesses at each `address` that have this single likely MISSING postal code value. 

Let's break this down into steps: 

### ❓ Part 1
Identify the single most common invalid postal code and assign it to `missing_postal_code`. Then create a `DataFrame`, `bus_missing`, to store only those businesses in `bus` that have `missing_postal_code` as their postal code.

**Hint**: All ZIP codes in the US are *positive* numbers

```python
missing_postal_code = ...
bus_missing = ...
```

```python
grader.check("q2ci")
```

### ❓ Part 2
Using `bus_missing`, find the number of businesses at each address (which would all share the same postal code). Specifically, `missing_zip_address_count` should store a `Series` with addresses as the indices and the counts as the values.

```python
missing_zip_address_count = ...
missing_zip_address_count.head()
```

```python
grader.check("q2cii")
```

<!-- BEGIN QUESTION -->

<br/>

--- 

## ❓ Question 2d

**(Written Response)** If we were to drop businesses with postal code values equal to `missing_postal_code`, what **specific types of businesses** would we be excluding? In other words, is there a commonality among businesses with missing postal codes? **Please respond in no more than 2 sentences.**

**Hint**: You may want to identify and Google the names of the businesses with missing postal codes. Feel free to reuse parts of your code from 2c to re-examine `bus_missing` but we will not be grading your code.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/>

--- 

## ❓ Question 2e

Examine the `invalid_zip_bus` `DataFrame` we computed in Question 2c and look at the businesses that DO NOT have the special MISSING ZIP code value. Some of the invalid postal codes are just the full 9-digit code rather than the first 5 digits. **Create a new column named `postal5` in the original `bus` `DataFrame`, which contains only the first 5 digits of the `postal_code` column.**

Then, for any of the `postal5` ZIP code entries that were not a valid San Francisco ZIP code (according to `valid_zip_codes`), the provided code will set the `postal5` value to `None`. 

**Hint:** You will find `str` accessors particularly useful. They allow you to use your usual Python string functions in tandem with a `DataFrame`. Refer to the [Pandas II course notes](https://ds100.org/course-notes/pandas-2/#adding-removing-and-modifying-columns) for examples.

**Do not modify the provided code! Simply add your own code in place of the ellipses.**


```python
bus['postal5'] = None
...

bus.loc[~bus['postal5'].isin(valid_zip_codes), 'postal5'] = None
# Checking the corrected postal5 column
bus.loc[invalid_zip_bus.index, ['bid', 'name', 'postal_code', 'postal5']]
```

```python
grader.check("q2e")
```

<br/>

---

<br/>

## ❓ Question 2f

Finally, use the `postal5` column to create a `DataFrame`, `bus_valid`, that only contains the rows of `bus` where a `postal5` zip code exists. You may find the `.isna()` function useful here.

```python
bus_valid = ...
bus_valid
```

```python
grader.check("q2f")
```

<br/>

---

<br/>

# 🔍 3: Investigate the Inspection Data

Let's now turn to the inspection `DataFrame`. Earlier, we found that `ins` has 4 columns named 
`iid`, `score`, `date`, and `type`.  In this section, we determine the granularity of `ins` and investigate the kinds of information provided for the inspections. 

Let's start by looking again at the first 5 rows of `ins` to see what we're working with.

```python
ins.head(5)
```

<br/>

---

## 🔍 Question 3a

The column `iid` probably corresponds to an inspection ID. Write an expression (i.e., a line of code) that evaluates to `True` or `False` based on whether all the inspection IDs are unique. Your code should compute the answer, i.e., don't just hard code `True` or `False`.

**Hint:** This is a very similar question to Question 1a.

```python
is_ins_iid_unique = ...
is_ins_iid_unique 
```

```python
grader.check("q3a")
```

<br/>

---

## 🔍 Question 3b

We want to extract `bid` from each row of the `ins` `DataFrame`. If we look carefully, the column `iid` of the `ins` `DataFrame` appears to be composed of two numbers, and the first number looks like a business ID.  

Create a new column called `bid` in the `ins` Dataframe containing just the business ID. You will want to use `ins['iid'].str` operations. (Python's in-built `split` method could come in use; read up on the documentation [here](https://pandas.pydata.org/docs/reference/api/pandas.Series.str.split.html)!) Also, be sure to convert the type of this column to `int`. 

**Hint**: Similar to an earlier problem where we used `astype("string")` to convert a column to a string, here you should use `astype` to convert the `bid` column into type `int`. **No Python `for` loops or list comprehensions are allowed.** This is on the honor system since our autograder isn't smart enough to check, but if you're using `for` loops or list comprehensions, you're doing the HW incorrectly. 

```python
ins['bid'] = ...
ins.head(5)
```

```python
grader.check("q3b")
```

<br/>

---

## 🔍 Question 3c

For this part, we're going to explore some new, somewhat strange syntax that we haven't seen in the lecture. Don't panic! If you're not sure what to do, try experimenting or Googling, and don't shy away from talking to other students or course staff.

For this problem, we'll use the time component of the inspection data.  All this information is given in the `date` column of the `ins` `DataFrame`. 

**No Python `for` loops or list comprehensions are allowed!**

### Part I

What is the type of the individual `ins['date']` entries? You may want to grab the very first entry and use the `type` function in Python. 

```python
ins_date_type = ...
ins_date_type
```

```python
grader.check("q3ci")
```

### 🔍 Part II
Rather than the type you discovered in Part 1, we want each entry in `pd.TimeStamp` format. You might expect that the usual way to convert something from its current type to `TimeStamp` would be to use `astype`. You can do that, but the more typical way is to use `pd.to_datetime` (documentation [here](https://pandas.pydata.org/docs/reference/api/pandas.to_datetime.html)). Using `pd.to_datetime`, create a new `ins['timestamp']` column containing `pd.Timestamp` objects. These will allow us to do date manipulation with much greater ease in parts III and IV. 

**Note:** You may encounter a `UserWarning` if you do not specify the date format when using pd.to_datetime. To resolve this, consider specifying the `format` using the string `'%m/%d/%Y %I:%M:%S %p'`. This format indicates that the datetime we are inputting to convert is in the form of `Month`/`Day`/`Year` `Hour` (12-hour clock)/`Minute`/`Second` `AM`/`PM`. Read [here](https://pandas.pydata.org/docs/reference/api/pandas.to_datetime.html) for more details.

```python
format_string = '%m/%d/%Y %I:%M:%S %p'
ins['timestamp'] = pd.to_datetime(...)
ins['timestamp']
```

```python
grader.check("q3cii")
```

### 🔍 Part III

What are the earliest and latest dates in our inspection data?  

**Hint**: you can use `min` and `max` on dates of the correct type.

```python
earliest_date = ...
latest_date = ...
print("Earliest Date:", earliest_date)
print("Latest Date:", latest_date)
```

```python
grader.check("q3ciii")
```

### 🔍 Part IV

We probably want to examine the inspections by year. Create an additional `ins['year']` column containing just the year of the inspection. Consider using `pd.Series.dt.year` to do this.

In case you're curious, the documentation for `TimeStamp` data can be found at [this link](https://pandas.pydata.org/docs/reference/api/pandas.Timestamp.html#pandas.Timestamp).

```python
ins['year'] = ...
```

```python
grader.check("q3civ")
```

<br/>

---

<br/>

# 🧑‍⚖️ 4: Benford's Law

[Benford’s Law](https://en.wikipedia.org/wiki/Benford%27s_law), also known as the Law of Anomalous Numbers or the First Digit Law, is a phenomenon seen in many real-world datasets. In datasets where Benford’s Law applies, there is an overrepresentation of small numbers (e.g., 1 and 2) in the leading digits of numeric entries and an underrepresentation of large numbers (e.g., 8 and 9). Benford’s Law is observed in river lengths, national populations, and has even been used to investigate [financial fraud](https://www.acfe.com/acfe-insights-blog/blog-detail?s=what-is-benfords-law-and-why-fraud-examiners-use-it)! More about Benford’s Law can be found [here](https://www.scientificamerican.com/article/what-is-benfords-law-why-this-unexpected-pattern-of-numbers-is-everywhere/).

<img src="https://static.scientificamerican.com/sciam/assets/File/BenfordsLaw_graphic_d%5B68%5D.jpg?w=1350" alt="Benford's Law Distribution" style="width: 600px">

Let’s investigate whether or not the **address** numbers in the `bus_valid` `DataFrame` follow Benford’s Law.

## 🧑‍⚖️ Question 4a

First, create a new column, `first_char` that contains the first character of the address string of each restaurant in `bus_valid`. **Again, do not use for loops or list comprehensions in your solution**.

```python
bus_valid["first_char"] = ...
bus_valid
```

```python
grader.check("q4a")
```

<br/>

---

## 🧑‍⚖️ Question 4b

Unfortunately, not all of the addresses in `bus_valid` start with a number. Create a new `DataFrame` `bus_digits` that only contains the rows of `bus_valid` where the restaurant's address starts with a digit between 1 and 9.

**Note: Do not include rows in which the first digit of the address is a 0.**

*Hint*: You may find the `s.str.isnumeric()` `Series` method useful here.

```python
help(pd.Series.str.isnumeric)
```

```python
bus_digits = ...
bus_digits
```

```python
grader.check("q4b")
```

<br/>

---

## 🧑‍⚖️ Question 4c

Plot a histogram that shows the frequency of each possible first digit among the addresses in `bus_digits`. Like you did in Homework 1, create bins of unit length that are centered at each of the digits and assign it to `bins_arr`.

**Note**: You only need to define `digits_arr` and `bins_arr`. The code necessary to plot the histogram has already been provided for you.

**Hint**: Make sure the `digits_array` has data type `int`. You can check the type of the object in a `Series` by looking at the `dtype` attribute. More information [here](https://pandas.pydata.org/docs/reference/api/pandas.Series.dtype.html). How can you change the content of the series to a specific data type?

**Hint:** To create bins of unit length that are centered at each of the digits, you might not want to start at the original starting and ending points. Where should you start instead?


```python
digits_arr = ...
bins_arr = ...

# DO NOT CHANGE BELOW
benford = np.array([0.301, 0.176, 0.125, 0.097, 0.079, 0.067, 0.058, 0.051, 0.046]) # Array of proportions under Benford's Law
plt.hist(digits_arr, bins = bins_arr, density = True, alpha = 0.5, color = "blue", label = "Empirical Dist") # Plot digits_arr, the empirical distribution
plt.bar(np.arange(1, 10), height = benford, width = 1, alpha = 0.5, color = "red", label = "Benford's Law") # Plot the theoretical distribution under Benford's Law
plt.xlabel("Digit")
plt.ylabel("Probability")
plt.title("Overlaid Plot of Empirical Distribution vs Benford's Law")
plt.legend();
```

```python
grader.check("q4c")
```

We have worked extensively with postal codes in this homework assignment. As prevalent and fascinating as it is, Benford's Law does not apply to *any* dataset. Benford's Law applies to datasets in which the numbers are *naturally generated*, which means that we will not be able to observe the same phenomenon in the leading digits of postal codes across the United States which are *manually assigned*. 

<img src="https://www.unitedstateszipcodes.org/images/zip-codes/zip-codes.png" alt="Leading Digit of ZIP Code across US" style="width: 600px">

<br/>

---

# 🏘️ 5: Search for the Perfect Restaurant

For this question, you work as a salesman. One of your clients is coming to San Francisco for a dinner meeting with you. You're in charge of picking the spot to eat. Your client wants to grab some food near Fisherman’s Wharf. They believe that the closer the restaurant is to the area the better the food will be. Your task is to use **coordinates** to choose the perfect restaurant for this dinner meeting to impress your important client.

## 🏘️ Question 5a

Create a new `DataFrame` `bus_coords`, which includes only the rows of `bus_valid` where both the latitude and longitude values are not missing. 

**Hint:** Missing values for latitude and longitude values are the *same* as the most frequent missing value you found in **Question 2c**!

**Hint:** You do not need to consider the scenario where a phone number is missing!

```python
bus_coords = ...
bus_coords
```

```python
grader.check("q5a")
```

<br/>

---

## 🏘️ Question 5b

To be close to Fisherman’s Wharf, find the **three restaurants** in the `bus_coords` `DataFrame` that are furthest north, breaking ties by prioritizing restaurants that are furthest east. Return the `name`, `address`, and `postal5` (columns in that order) of the three restaurants as a `DataFrame` and assign it to `top3`. Here is an [article](https://www.britannica.com/science/latitude) that succinctly explains latitude and longitude.

**Hint:** Feel free to reference the [sort_values documentation](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.sort_values.html) to see how you can sort by multiple values.

```python
# SCRATCH CELL
# Feel free to do your rough work here
# Do not add a cell between your solution and the grader cell
```

```python
top3 = ...
top3
```

```python
grader.check("q5b")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Bart says congratulations! You have finished Homework 2A! ##

<img src = "bart.jpeg" width = "250">

### Course Content Feedback

If you have any feedback about this assignment or about any of our other weekly, weekly assignments, lectures, or discussions, please fill out the [Course Content Feedback Form](https://forms.gle/9MpRsWt51ktcwugW9). Your input is valuable in helping us improve the quality and relevance of our content to better meet your needs and expectations!

### Submission Instructions

Below, you will see a cell. Running this cell will automatically generate a zip file with your autograded answers. Once you submit this file to the HW 2A Coding assignment on Pensieve, Pensieve will automatically submit a PDF file with your written answers to the HW 2A Written assignment. 

**Important**: Please check that your written responses were generated and submitted correctly to the HW 2A Written Assignment. 

**You are responsible for ensuring your submission follows our requirements and that the PDF for HW 2A written answers was generated/submitted correctly. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1a (points: 1)
```
>>> assert bool(is_bid_unique or ~is_bid_unique)
```


Tests for q1b (points: 2)
```
>>> assert len(top_names) >= 6
>>> assert len(top_addresses) >= 6
```
```
>>> assert top_names[-1] == "Specialty's Cafe & Bakery" or top_names[-1] == 'Mixt Greens/Mixt'
>>> assert top_addresses[-1] == '24 Willie Mays Plaza'
```


Tests for q1c (points: 1)
```
>>> assert q1c.upper() in set(['A', 'B', 'C', 'D'])
```


Tests for q2a (points: 1)
```
>>> assert type(zip_counts) == pd.Series
```
```
>>> assert zip_counts.shape[0] > 50
```


Tests for q2bi (points: 1)
```
>>> assert type(valid_zip_codes) == pd.Series
```
```
>>> assert valid_zip_codes[0] == 94102
```
```
>>> assert bool(valid_zip_codes.dtype == np.int64)
```
```
>>> assert bool(np.all(valid_zip_codes.sort_values() == pd.read_json('data/sf_zipcodes.json')['zip_codes'].sort_values()))
```


Tests for q2bii (points: 1)
```
>>> assert type(invalid_zip_bus) == pd.DataFrame
>>> assert len(invalid_zip_bus) > 200
```


Tests for q2ci (points: 1)
```
>>> assert type(missing_postal_code) == str
```
```
>>> assert bus_missing.shape[0] == 194
```


Tests for q2cii (points: 1)
```
>>> assert type(missing_zip_address_count) == pd.Series
```
```
>>> assert len(missing_zip_address_count) == 135
>>> assert bool(missing_zip_address_count['3914 Judah St'] == 1)
```


Tests for q2e (points: 2)
```
>>> assert 'postal5' in bus.columns
>>> assert bool((bus['postal5'].str.len() != 5).sum() == 221)
```
```
>>> assert bus['postal5'].isin(valid_zip_codes).sum() == 6032
>>> assert bool(bus['postal5'].isna().sum() == 221)
```


Tests for q2f (points: 1)
```
>>> assert bool(bus_valid.shape[0] > 6000)
```


Tests for q3a (points: 1)
```
>>> assert type(is_ins_iid_unique) == bool or type(is_ins_iid_unique) == np.bool_
```


Tests for q3b (points: 2)
```
>>> assert 'bid' in ins.columns
```
```
>>> assert ins['bid'].dtype == int
```
```
>>> assert len(ins[ins['score'] > 0]['bid'].unique()) > 5000
```


Tests for q3ci (points: 1)
```
>>> assert type(ins_date_type) == type
```


Tests for q3cii (points: 1)
```
>>> assert type(ins['timestamp'][1]) == pd.Timestamp
```


Tests for q3ciii (points: 1)
```
>>> assert type(earliest_date) == pd.Timestamp
```
```
>>> assert type(latest_date) == pd.Timestamp
```


Tests for q3civ (points: 1)
```
>>> assert 'year' in ins.columns
```


Tests for q4a (points: 1)
```
>>> assert len(bus_valid.columns) == 11 and 'first_char' in bus_valid.columns
```
```
>>> assert bool(bus_valid.shape == (6032, 11))
```


Tests for q4b (points: 1)
```
>>> assert 5000 < bus_digits.shape[0] < 6000
```


Tests for q4c (points: 1)
```
>>> assert bool(len(bins_arr) == 10)
```


Tests for q5a (points: 2)
```
>>> assert 1000 < bus_coords.shape[0] < 3000
```


Tests for q5b (points: 2)
```
>>> assert bool(top3.shape[0] == 3)
```
```
>>> assert bool(type(top3) == type(pd.DataFrame()))
```
